```python
# Initialize Otter
import otter
grader = otter.Notebook("lab3.ipynb")
```

```python
import numpy as np
import matplotlib.pyplot as plt
import pickle
import scipy.stats as stats
```

---

<h1><center>SDSE Lab 3 <br><br> K-means clustering </center></h1>

---

# 1) Load the data
We load the dataset from the file `'data2d.pickle'`. The data matrix is called `X`. Save the number of samples (rows) as `N` and the number of features (columns) as `D`.

```python
with open('...','...') as f:    # TODO
    X = pickle.  ...    # TODO
N = ...    # TODO
D = ...    # TODO
```

```python
grader.check("q1")
```

# 2) Scatter plot
Create a scatter plot with the first column of `X` along the x-axis and the second column along the y-axis. Use `plt.axis('equal')` to force the two axes to have the same scale. This is important for correct visualization of the Euclidean distances between points. Your figure should look like the one below.

Food for thought: How many clusters do you see?

<img src="2.png" />

```python
fig, ax = plt.subplots(figsize=(5,5))
...  # TODO
```

# 3) Random placement of centroids
Write a function that takes an integer `K` and the data matrix `X` as inputs, and returns `K` centroids randomly placed within the box spanned by the columns of `X`. The shape of the return value should be `(K,D)`. 

The template also takes a third optional parameter. This is used by the autograder to ensure repeatable results. 

**Note**: You should use `stats.uniform.rvs(loc=..., scale=..., size=...)` to generate random centroid placements. 

```python
def initialize_centroids(K,X,random_seed=None):
    if random_seed is not None:
        np.random.seed(seed=random_seed)
    D = X.shape[1]
    C = np.empty((K,D))
    ...    # TODO
    return C
```

```python
C4 = initialize_centroids(4,X,random_seed=234)
C8 = initialize_centroids(8,X,random_seed=154)
```

```python
grader.check("q3")
```

# 4) Compute the assignment for given centroids
Write a function that takes the data matrix `X` (shape `(N,D)`) and centroid locations `C` (shape `(K,D)`) as inputs, and returns an assignment matrix `gamma`. `gamma` should be a numpy array with shape `(N,K)`. Each row of `gamma` should contain `K-1` `False` values and one `True` value in the position corresponding to the centroid that is nearest to that data point (Euclidean distance). Hint: `np.argmin`.

```python
def assign_samples_to_nearest_centroid(X,C):
    N = ...  # TODO
    K = ...  # TODO
    gamma = np.zeros((N,K),dtype=bool)  # Initialized to all False
    ...  # TODO
    return gamma
```

```python
gamma4 = assign_samples_to_nearest_centroid(X,C4)
gamma8 = assign_samples_to_nearest_centroid(X,C8)
```

```python
grader.check("q4")
```

# 5) Plot data, centroids, assignment
Write a function that takes as input the data matrix `X`, a centroids matrix `C`, and an assignment matrix `gamma`, and creates a scatter plot similar to the one shown here. The centroids should be indicated with a large dot, and the data samples should be color-coded according to their cluster membership. Your plot may not look identical to this one due to the random placement of the centroids.

<img src="5.png" />


```python
def plot_scatter(X, C, gamma):
    K = ...  # TODO
    fig, ax = plt.subplots(figsize=(5,5))
    for k in range(K):
        ax.plot(...,...,'.', markersize=10)  # plot data assigned to cluster k
    ax.plot(...,..., 'o', markersize=10)  # plot the centroids
    ax.axis('equal')
```

```python
plot_scatter(X, C4, gamma4)
```

# 6) Compute the centroids for a given assignment
Write a function that takes the assignment `gamma` and the data matrix `X` and returns centroids placed at the mean of each cluster. The return value should be a numpy array with shape `(K,D)`.

```python
def place_centroids(gamma, X):
    K = gamma.shape[1]
    D = X.shape[1]
    C = np.empty((K,D))
    ...    # TODO
    return C
```

```python
newC4 = place_centroids(gamma4, X)
newC8 = place_centroids(gamma8, X)
```

```python
grader.check("q6")
```

# 7) Code K-means
Now we will implement the K-means algorithm using the methods that we've coded so far. Implement the K-means algorithm described in class. 

The function `run_kmeans` takes the dataset `X` and initial centroids `C` as input. It returns the final placement of the centroids `C`, and the assignment matrix `gamma`. There are two stopping conditions:
+ Stop and return `C=None` and `gamma=None` if any of the clusters becomes empty. 
+ Stop if the assignment matrix does not change from one iteration to the next. 

```python
def run_kmeans(X,C):

    done = False
    first = True
        
    while not done:
        
        # compute assignment (E step)
        gamma = ...
        
        # place centroids (M step)
        C = ...
        
        # Break if any of the clusters are empty
        # Return C=None and gamma=None
        if ...:
            ...
            
        # Break if there is no change in the assignment
        if ...:
            ...
            
        # update gamma
        gammaold = gamma
        first = False
            
    return C, gamma
```

```python
optC4, optgamma4 = run_kmeans(X,C4)
optC8, optgamma8 = run_kmeans(X,C8)
```

```python
grader.check("q7")
```

# 8) Evaluate the cost function 
Write a function that evaluates the cost function for a given dataset `X`, centroid placement `C`, and assignment matrix `gamma`. The cost function is the sum of the squared distances from each point (row in `X`) to its assigned centroid (row in `C`). The function should return `np.nan` if `C` is `None`.

```python
def eval_cost(X,C,gamma):
    if C is None:
        return np.nan
    K = ...   # TODO
    cost = ...   # TODO
    return cost
```

```python
J4 = eval_cost(X,optC4,optgamma4)
J10 = eval_cost(X,optC8,optgamma8)
```

```python
grader.check("q8")
```

# 9) Ensemble run
Write a function that runs K-means multiple times, each with a different initial centroid placement. These initial placements should be chosen with `initialize_centroids(K,X)`. The function returns `best_run`, which is a dictionary containing the centroids (`C`), cost (`cost`), and assignment matrix (`gamma`) for the best of the ensemble of runs. 

Note that the function also takes `random_seed` as a fourth argument. This is used by the autograder to ensure repeatability of the results. You should **not** pass `random_seed` to `initialize_centroids`.

```python
def ensemble_run(K,X,ensemblesize,random_seed=None):
    if random_seed is not None:
        np.random.seed(seed=random_seed)
    best_run = {'C':None,'gamma':None,'cost':None}
    for e in range(ensemblesize):
        
        # Run Kmeans with K randomly initialized centroids
        # If an empty cluster is obtained, then repeat, and keep on repeating 
        # until the result has no empty clusters
        ...   # TODO

        # Evaluate the cost function 
        ...   # TODO

        # Check whether this is the best result obtained so far, 
        # and if so, store it in best_run.
        ...   # TODO
  
    return best_run
```

```python
best_run_K4 = ensemble_run(4,X,ensemblesize=100,random_seed=5423)

# sort the result for the autograder to work
ind = np.argsort(np.linalg.norm(best_run_K4['C'],axis=1))
best_run_K4['C'] = best_run_K4['C'][ind,:]
best_run_K4['gamma'] = best_run_K4['gamma'][:,ind]
```

```python
grader.check("q9")
```

# 10) Iteration over K
Complete the function below that takes the data matrix `X`, an ensemble size `ensemblesize`, and a maximum number of clusters to consider (`maxK`).  The function should run `ensemble_run` for all values of `K` from 1 to `maxK`, and return a numpy array of length `maxK` with the best case cost for each value of `K`. 

A random seed is passed in to the function for the autograder. Please pass this seed on to the `ensemble_run` function.

```python
def K_iteration(X, ensemblesize, maxK,random_seed=None):
    if random_seed is not None:
        np.random.seed(seed=random_seed)
    best_cost = np.empty(maxK)
    for ind, K in enumerate(range(...)):   # TODO
        print(K)
        ...   # TODO
    return best_cost
```

```python
ensemblesize=100
maxK=10
best_cost = K_iteration(X, ensemblesize, maxK,random_seed=452)
```

```python
grader.check("q10")
```

# 11) Selecting K

Running the cell below will produce a plot showing the performance of the algorithm for different values of `K` ranging from 1 to 10 (calculated in the previous part). The top row shows the intertia cost function, which naturally decreases monotonically with increasing `K`. The bottom plot shows the percentage improvement from one value of `K` to the next. What is the best value of `K` to choose according to these plots? 

```python
fig, ax = plt.subplots(figsize=(8,5),nrows=2,sharex=True)

ax[0].plot(range(1,maxK+1),best_cost,
         linewidth=3,
         marker='o',
         markersize=8)
ax[0].grid()
ax[0].set_ylabel('Inertia')

ax[1].plot(range(2,maxK+1),100*np.abs(np.diff(best_cost)) / best_cost[1:],
         linewidth=3,
         marker='o',
         markersize=8)
ax[1].grid()
ax[1].set_ylabel('Inertia improvement')
ax[1].set_xlabel('K')
```

```python
Kideal = ...   # TODO
```

```python
grader.check("q11")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> N==100 and D==2
True
```


Tests for q10 (points: None)
```
>>> np.all(np.isclose(best_cost,np.array([2895.67902369,  981.02477219,  616.56731334,  265.58230558,
...         223.91615237,  192.59481366,  162.47747163,  142.77422271,
...         126.07926529,  112.32417365])))
True
```


Tests for q11 (points: None)
```
>>> Kideal == 4
True
```


Tests for q3 (points: None)
```
>>> np.all(np.isclose(C4,np.array([[-10.8395711 ,  -1.75265235],
...        [ -8.05364821,   5.32904513],
...        [  0.15658959,   7.95416692],
...        [  0.87733571,   4.40057558]]),1e-3))
True
```
```
>>> np.all(np.isclose(C8,np.array([[-4.89188176,  7.71937049],
...        [-3.35957754,  3.81687063],
...        [-2.67324835, -3.41980139],
...        [ 1.06844648,  7.32753667],
...        [-0.41543367,  0.92345392],
...        [-0.89173646,  2.72850835],
...        [-2.28524226,  5.16074095],
...        [-3.77139276,  8.13414145]]),1e-3))
True
```


Tests for q4 (points: None)
```
>>> np.all(np.where(gamma4[:,0])[0] == np.array([ 1,  3,  5, 10, 13, 18, 22, 25, 27, 31, 35, 38, 39, 41, 42, 47, 53,
...        55, 65, 67, 70, 71, 77, 86, 94, 98]))
True
```
```
>>> np.all(np.where(gamma4[:,1])[0] == np.array([ 2,  6,  9, 12, 17, 20, 23, 24, 26, 28, 32, 33, 40, 43, 44, 46, 48,
...        51, 56, 57, 58, 60, 61, 64, 68, 69, 72, 75, 76, 78, 80, 81, 84, 89,
...        90, 91, 95, 96]))
True
```
```
>>> np.all(np.where(gamma8[:,0])[0] == np.array([ 6,  9, 12, 17, 20, 23, 24, 28, 32, 46, 48, 51, 56, 57, 58, 60, 61,
...        64, 68, 69, 72, 75, 80, 89, 90, 91, 96]))
True
```
```
>>> np.all(np.where(gamma8[:,1])[0] == np.array([ 2, 11, 18, 25, 26, 29, 33, 40, 43, 54, 62, 65, 67, 76, 78, 81, 84, 92, 94, 95]))
True
```


Tests for q6 (points: None)
```
>>> np.all(np.isclose(newC4,np.array([[-6.13653527, -1.60252754],
...        [-7.39643996,  6.51996032],
...        [-2.9212051 ,  6.22055446],
...        [-1.56396103, -1.18061685]]),1e-3))
True
```
```
>>> np.all(np.isclose(newC8[:3,:],np.array([[-8.43821088,  7.2752506 ],
...        [-5.17362655,  3.0960672 ],
...        [-3.30674942, -2.56457236]]),1e-3))
True
```
```
>>> np.all(np.where(np.isnan(newC8[:,0]))[0]==np.array([3, 4, 5, 7]))
True
```


Tests for q7 (points: None)
```
>>> np.all(np.isclose(optC4,np.array([[-6.37194325, -1.42726545],
...         [-8.73054661,  7.38064932],
...         [-3.80111472,  4.63859539],
...         [-1.33733821, -3.09356399]]),1e-3))
True
```
```
>>> (optC8 is None) and (optgamma8 is None)
True
```


Tests for q8 (points: None)
```
>>> np.isclose(J4,265.5823055768351,1e-3)
True
```
```
>>> np.isnan(J10)
True
```


Tests for q9 (points: None)
```
>>> np.all(np.isclose(best_run_K4['C'][np.argsort(np.linalg.norm(best_run_K4['C'],axis=1)),:],
...        np.array([[-1.33733821, -3.09356399],
...        [-3.80111472,  4.63859539],
...        [-6.37194325, -1.42726545],
...        [-8.73054661,  7.38064932]]),1e-3))
True
```
```
>>> np.all(np.where(best_run_K4['gamma'][:,3])[0]==np.array([ 6,  9, 12, 17, 20, 23, 24, 28, 32, 48, 51, 56, 57, 58, 60, 61, 64,
...        69, 72, 75, 80, 89, 90, 91, 96]))
True
```
```
>>> np.isclose(best_run_K4['cost'],265.5823055768351,1e-3)
True
```
