```python
# Initialize Otter
import otter
grader = otter.Notebook("hw5.ipynb")
```

---

<h1><center>SDSE Homework 5 <br><br> Neural networks for time series forecasting </center></h1>

---

In this homework we will build a model to forecast the intensity of traffic flow on a freeway over the coming 24 hours, based on measurements of traffic flow intensity over the past six hours. The data was obtained from Caltrans' Performance Measurement System ([PeMS](https://pems.dot.ca.gov/)). This site archives 5-minute average flow and speed measurements taken from sensors embedded in the pavement, at approximately 1-mile intervals, on freeways across the state. We will focus on measurements from a particular sensor, taken between January 2018 and October 2019. 

We will compare three neural network architectures: a dense neural network with multiple input and output channels, a simple recurrent neural network, and an LSTM network. We will use Keras and Tensorflow to build and train these models. Use `pip` to install these packages in your Python environment.

``` 
pip install tensorflow
pip install keras
```

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.preprocessing import StandardScaler
import tensorflow as tf
import keras
from tensorflow.keras.models import Sequential
from tensorflow.keras.layers import Dense, LSTM, SimpleRNN, Flatten
keras.utils.set_random_seed(812)
tf.config.experimental.enable_op_determinism()
```

# 0. Load the data

Use `pd.read_csv` to load the CSV file `'traffic_flows.csv'` data into a pandas DataFrame called `rawdata`.

```python
rawdata = ...
```

```python
grader.check("q0")
```

Notice that we have a total of 13,980 data samples, ranging from January 1st, 2018 to October 26th, 2019. The following table tells us that the `DateAndTime` column is string-valued, while the `Total Flow` column contains numbers.

```python
rawdata.info()
```

## 1. Convert the DateAndTime column to type Timestamp

To enable pandas' time series features, we must convert the `DateAndTime` information from a strings into `Timestamp` objects. Use pandas' [`to_datetime`](https://pandas.pydata.org/docs/reference/api/pandas.to_datetime.html) method to do this. Pass the `DateAndTime` column of `rawdata` into this method, and store the result in a new column called 'DateAndTimeTS'

```python
rawdata['DateAndTimeTS'] = ...
```

```python
rawdata.head()
```

```python
grader.check("q1")
```

## 2. Use the `DateAndTimeTS` as the index

The [`set_index`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.set_index.html) method of `rawdata` returns a new table with the given column as its index. Pass the string 'DataAndTimeTS' into this method and store the result in `data1`.

```python
data1 = ...
```

```python
data1.head()
```

```python
grader.check("q2")
```

# 3. Drop the `DateAndTime` column

The `DateAndTime` column is now superfluous. Use the [`drop`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.drop.html) method applied to `data1` with the `columns` argument to remove this column. Store the result in `data2`.

```python
data2 = ...
```

```python
data2.head()
```

```python
grader.check("q3")
```

Pandas attaches matplotlib plotting functions to its DataFrame objects. Hence we can call `plot` directly from the DataFrame.

```python
data2.plot(figsize=(10,4))
```

As a side note, you can add pan/zoom functionality to your plots by installing the `ipympl` package (`pip install ipympl`) and adding the line `%matplotlib widget` to your code:

``` python 
%matplotlib widget
plt.figure(figsize=(10,4))
plt.plot(data2,'.-')
plt.grid()
```
Here is what we see if we then zoom in to the plot:

<img src="f1.png" />

We can also pass in the x and y limits to pandas' plot method. Notice below that the x-limits can be conveniently specified as dates (this is one of those aformentioned feature).

```python
data2.plot(figsize=(10,4),
           xlim=('2018-02-16','2018-03-16'),
           style='.-')
```

There are indeed quite a few days missing in February and March 2018. If these gaps were small, we might consider filling them with linear interpolation -- effectively filling in the above plot with round markers along the long lines. However this would likely confuse the model-training procedure. Instead we will simply ignore these gaps. The model-training procedure will still be confused by the suddem jumps in the flow, however perhaps less than with interpolation. 

# 4. Training and testing datasets

We will use 80\% of the data for training and 20\% for testing. The dataset has 13,980 hourly samples. Since we are building a time-series forecasting model, it is important that the test dataset follow the training dataset in time. Create a training dataset (`data_train`) by selecting the first 11,184 samples (`ind_split`) from `data2`, and a testing dataset with the remaining samples. 

**Hint**: [`data2.iloc`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.iloc.html)

```python
Nraw = data2.shape[0]
train_portion = 0.8
ind_split = int(0.8*Nraw)
print(Nraw, ind_split)
```

```python
data_train = data2.iloc[...]
data_test = data2.iloc[...]
```

```python
grader.check("q4")
```

```python
plt.figure(figsize=(10,4))
plt.plot(data_train,label='training data')
plt.plot(data_test,label='test data')
plt.legend()
```

# 5. Normalize the data

The performance of neural network models is typically enhanced by normalizing the inputs and outputs. We will use scikit-learn's `StandardScaler` to do this. 

When preprocessing data, such as we are doing now, it is important that the test data be treated identically to the training data. This means that the same scaling factors should be applied to the test data samples as were applied to the training data. This was achieved in lab 5 with scikit-learn's `Pipeline` object. Here we will achieve the same by scaling the test data with `StandardScaler` that was fit using the training data.

**Hint** The `fit` and `transform` methods of the `StandardScaler` must receive a matrix-like object -- that is, one with two dimensions. This means that we must add an empty dimension to our one-dimensional flow vector. To acheive this, you can use the `reshape` method as follows:

`data_train['Flow'].values.reshape(-1, 1)`

```python
scaler = StandardScaler()
scaler.fit(...)

data_train_scaled = scaler.transform(...)
data_test_scaled = scaler.transform(...)
```

```python
grader.check("q5")
```

# 6. Organize the data into past and future snippets of flow values.

Write a function that recevies a scaled dataset `d` (`data_train_scaled` or `data_test_scaled`) and creates inputs and output matrices `X` and `y` suitable for model training.

The figure below provides an illustration. Each sample $(x_i,y_i)$ consists of `Np=6` hours of "past" data in $x_i$ and `Nf=24` hours of "future" data in $y_i$. If the input sequence `d` has length `L`, then there will be `N=L-Np-Nf` such samples. Hence the dimensions of `X` and `y` should be `(N,Nh)` and `(N,Nf)` respectively.

<img src='f2.png' />

```python
Np = 6
Nf = 24

def make_samples(d):

    L = d.shape[0]
    N = ...

    X = np.empty(...)
    y = np.empty(...)

    for i in range(N):
        X[i,:] = d[...,0]
        y[i,:] = d[...,0]

    return X, y
```

```python
Xtrain, ytrain = make_samples(data_train_scaled)
Xtest, ytest = make_samples(data_test_scaled)

Xtest.shape, ytest.shape
```

```python
grader.check("q6")
```

```python
i = 40

plt.figure(figsize=(10,4))
plt.plot(np.arange(0,Np),Xtrain[i],'ro-',linewidth=2,label='x_i')
plt.plot(np.arange(Np,Np+Nf),ytrain[i],'go-',linewidth=2,label='y_i')
plt.legend()
plt.grid(':')
```

# 7. Add an empty dimension to the training and test datasets to satisfy a Tensorflow requirement

```python
Xtrain_e = Xtrain[:,:,np.newaxis]
ytrain_e = ...
Xtest_e = ...
ytest_e = ...
```

```python
grader.check("q7")
```

# 8. Train a dense neural network

A schematic of our model is shown below.

<img src='f3.png'>

The model has a single dense hidden layer with 5 neurons (a.k.a. "units") and a `relu` activation function. The output layer has `Nf=24` units -- one  for each of the outputs of the model. Keras requires that we add a `Flatten()` layer to process the input (I'm not sure why). The final model code looks like this:

``` python
model_dense = Sequential([
    Flatten(),
    Dense(5, activation="relu"),
    Dense(Nf)
])
```

Enter this code in the cell below.

```python
model_dense = ...
```

```python
grader.check("q8")
```

Next we compile and train this model. The compilation step was not required in scikit-learn. It is used in Keras to configure the model object according to the hardware characteristics of the computing environment. For example, it will prepare the model to be trained using GPUs if you have configured Keras to do this and have the appropriate hardware. 

The `compile` method must receive the chosen loss function to be minimized, and the numerical optimization method to use. We will use minimize the mean squared error using the `adam` optimizer (an improved version of stochastic gradient descent).

```python
model_dense.compile(loss="mse",optimizer="adam")
```

Finally, we train the model using the `fit` function. Similarly to scikit-learn, Keras' `fit` function takes the training data as input. It also requires the number of epochs (iterations through the training dataset) to run. We will train all of our neural network models for 25 epochs. The method returns a logging object that contains the evolution of the loss over the epochs. Save this to `history_dense`

```python
train_epochs = 25
history_dense = model_dense.fit(Xtrain_e, ytrain_e,epochs=train_epochs)
```

# 9. Train a simple recurrent neural network

Next we will build a neural network with simple recurrent units. The schematic is identical to the one shown in section 8, except that now the units in the hidden layer will be recurrent neurons. In Keras, this is specified as a sequential model with two layers. The `Flatten` layer of the dense network is no longer required. Instead, the first layer is of type `SimpleRNN`. The constructor for `SimpleRNN` takes the number of recurrent units (5) and the shape of the input vector (`input_shape=(Np,1)`). The second layer is the same output layer as was used for the dense network.

```python 
model_simprnn = Sequential([
    SimpleRNN(5,input_shape=(Np,1)),
    Dense(Nf)
])
```
Create this model, then compile and train it with the same arguments as in part 8.


```python
model_simprnn = ...
model_simprnn.compile(...)
history_simprnn = model_simprnn.fit(...)
```

```python
grader.check("q9")
```

# 10. Train an LSTM neural network

Finally we will create a neural network with LSTM units. The code is identical to the Simple RNN case, except that instad of a `SimpleRNN` layer, we use an `LSTM` layer. Create, compile, and train this model. Save the model to `model_lstm` and its training history to `history_lstm`.

```python
model_lstm = ...
model_lstm.compile(...)
history_lstm = model_lstm.fit(...)
```

```python
grader.check("q10")
```

# 11. Forecast

Use the `predict` method of the model objects to generate forecasts for the test data. The input to this function is the "extended" version of the testing data `Xtest_e`. 

```python
yhat_dense   = model_dense.predict(Xtest_e)
yhat_simprnn = ...
yhat_lstm    = ...
```

```python
grader.check("q11")
```

# 12. Plots

### 12.1 Training histories

The loss after each of the training epochs is contained in the history object, and is plotted below. Notice that the LSTM model concludes the training with the lowest loss value, while the dense and Simple RNN models have similar values of final training loss. 

```python
plt.figure(figsize=(10,4))
plt.plot(history_dense.history['loss'], 'bo-',label='Dense')
plt.plot(history_simprnn.history['loss'], 'mo-',label='Simple RNN')
plt.plot(history_lstm.history['loss'], 'co-',label='LSTM')
plt.legend()
plt.grid(':')
plt.xlabel('Epochs')
plt.ylabel('Training MSE')
```

### 12.2. A sample prediction

To plot the forecast, we will first use the `scaler` object to return the data to its original scale (in vehicles per hour, or vph).

```python
Xtest_vph = scaler.inverse_transform(Xtest)
ytest_vph = scaler.inverse_transform(ytest)
yhat_dense_vph = scaler.inverse_transform(yhat_dense)
yhat_simprnn_vph = scaler.inverse_transform(yhat_simprnn)
yhat_lstm_vph = scaler.inverse_transform(yhat_lstm)
```

```python
i=500
tx = np.arange(-Np,0)+1
ty = np.arange(Nf)+1

fig, ax = plt.subplots(figsize=(10,5)) 
ax.plot(tx,Xtest_vph[i],'ro-', lw = 3,label='past 6 hours') 
ax.plot(ty,ytest_vph[i] ,'go-', lw = 3,label='future 24 hours') 
ax.plot(ty,yhat_dense_vph[i],'bo-', markersize=3, lw = 1, label='Dense') 
ax.plot(ty,yhat_simprnn_vph[i],'mo-', markersize=3, lw = 1, label='Simple RNN') 
ax.plot(ty,yhat_lstm_vph[i],'co-', markersize=3, lw =1, label='LSTM') 
ax.grid(':')
ax.legend()
```

# 13. Compute performance

We measure the performance of the models with their mean absolute errors. Write a function that computes the MAE between two arrays. Apply it to the unscaled versions of the three model predictions and the test outputs.

```python
def compute_mae(yhat,ytest):
    return ...
```

```python
mae_dense = compute_mae(yhat_dense,ytest_vph)
mae_simprnn = compute_mae(yhat_simprnn_vph,ytest_vph)
mae_lstm = compute_mae(yhat_lstm_vph,ytest_vph)

mae_dense, mae_simprnn, mae_lstm
```

```python
grader.check("q13")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q0 (points: None)
```
>>> (rawdata is not None) and rawdata.shape[0]==13980
True
```


Tests for q1 (points: None)
```
>>> ('DateAndTimeTS' in rawdata.columns) and (type(rawdata['DateAndTimeTS'][0])==pd._libs.tslibs.timestamps.Timestamp)
True
```


Tests for q10 (points: None)
```
>>> model_lstm.get_config()['layers'][0]['class_name']=='InputLayer' and \
... model_lstm.get_config()['layers'][0]['config']['batch_input_shape']==(None, 6, 1) and \
... model_lstm.get_config()['layers'][1]['class_name']=='LSTM' and \
... model_lstm.get_config()['layers'][1]['config']['batch_input_shape']==(None, 6, 1)
True
```


Tests for q11 (points: None)
```
>>> yhat_dense.shape==(2766, 24) and yhat_simprnn.shape==(2766, 24) and yhat_lstm.shape==(2766, 24)
True
```
```
>>> np.all(np.isclose(yhat_dense[:5,:5],np.array([[0.4388529 , 0.64404047, 0.6783319 , 0.73585904, 0.84136415],
...        [0.97280717, 0.9306428 , 0.79587066, 0.78927994, 0.8791541 ],
...        [1.013487  , 0.99908984, 0.92824966, 0.9273282 , 0.9662    ],
...        [0.9781668 , 1.0277343 , 1.0271584 , 0.99955845, 0.94756854],
...        [1.0915511 , 1.1103691 , 1.1362342 , 1.1358428 , 1.0763679 ]]),1e-2))
False
```
```
>>> np.all(np.isclose(yhat_simprnn[:5,:5],np.array([[0.72212905, 0.96939445, 1.0435158 , 1.0437351 , 1.0022954 ],
...        [0.89051837, 0.9253676 , 0.8325755 , 0.77615154, 0.7460695 ],
...        [0.9098348 , 0.8963374 , 0.82915664, 0.8337238 , 0.8794229 ],
...        [0.85661846, 0.83173513, 0.811392  , 0.8633232 , 0.9311729 ],
...        [0.86052257, 0.83580565, 0.83853936, 0.87950665, 0.9114979 ]]),1e-2))
False
```
```
>>> np.all(np.isclose(yhat_lstm[:5,:5],np.array([[0.51328224, 0.8539859 , 1.006242  , 1.0394508 , 1.0034708 ],
...        [0.76915145, 1.0135615 , 1.0681678 , 1.0212207 , 0.92616224],
...        [0.91571105, 1.0441518 , 1.0261366 , 0.9555544 , 0.8570435 ],
...        [1.0291822 , 1.0403317 , 0.9755738 , 0.92160684, 0.85469663],
...        [0.88564634, 0.81100476, 0.7573972 , 0.7753442 , 0.8072547 ]]),1e-2))
False
```


Tests for q13 (points: None)
```
>>> np.isclose(mae_dense,1.0000951034154313,1e-2)
True
```
```
>>> np.isclose(mae_simprnn,0.2985688091508766,1e-2)
False
```
```
>>> np.isclose(mae_lstm,0.28611520439895655,1e-2)
False
```


Tests for q2 (points: None)
```
>>> (data1 is not None) and (data1.index.name=='DateAndTimeTS') and (type(data1.index[0])==pd._libs.tslibs.timestamps.Timestamp)
True
```


Tests for q3 (points: None)
```
>>> (len(data2.columns)==1) and (data2.columns[0]=='Flow')
True
```


Tests for q4 (points: None)
```
>>> data_train.shape==(11184, 1) and data_test.shape==(2796, 1) and data_train.index[0]==pd.Timestamp('2018-01-01 00:00:00') and data_train.index[-1]==pd.Timestamp('2019-06-26 00:00:00') and data_test.index[0]==pd.Timestamp('2019-06-26 01:00:00') and data_test.index[-1]==pd.Timestamp('2019-10-26 06:00:00') 
True
```


Tests for q5 (points: None)
```
>>> np.isclose(scaler.mean_[0],7218.13796495,1e-1) and np.isclose(scaler.scale_[0],2959.82974674,1e-1)
True
```
```
>>> data_train_scaled.shape==(11184, 1) and np.all(np.isclose(data_train_scaled[:10,0],np.array([-1.2099135 , -1.10247489, -1.26836281, -1.48932146, -1.54033791, -1.59203007, -1.43560216, -1.32444711, -1.26971424, -0.96496698]),1e-3))
True
```
```
>>> data_test_scaled.shape==(2796, 1) and np.all(np.isclose(data_test_scaled[:10,0],np.array([-1.57479935, -1.58662436, -1.53358076, -1.35721927, -0.83624336,0.17732845,  0.86520586,  1.10339523,  0.89899158,  0.75878082]),1e-3))
True
```


Tests for q6 (points: None)
```
>>> Xtrain.shape==(11154, 6) and ytrain.shape==(11154, 24) and Xtest.shape==(2766, 6) and ytest.shape==(2766, 24)
True
```
```
>>> np.all(np.isclose(Xtrain[:3,:],np.array([[-1.2099135 , -1.10247489, -1.26836281, -1.48932146, -1.54033791,-1.59203007],
...        [-1.10247489, -1.26836281, -1.48932146, -1.54033791, -1.59203007,-1.43560216],
...        [-1.26836281, -1.48932146, -1.54033791, -1.59203007, -1.43560216,-1.32444711]]),1e-3))
True
```
```
>>> np.all(np.isclose(ytrain[:3,:6],np.array([[-1.43560216, -1.32444711, -1.26971424, -0.96496698, -0.74434618,-0.36087818],
...        [-1.32444711, -1.26971424, -0.96496698, -0.74434618, -0.36087818,-0.04160306],
...        [-1.26971424, -0.96496698, -0.74434618, -0.36087818, -0.04160306,-0.02268305]]),1e-3))
True
```


Tests for q7 (points: None)
```
>>> Xtrain_e.shape==(11154, 6, 1) and ytrain_e.shape==(11154, 24, 1) and Xtest_e.shape==(2766, 6, 1) and ytest_e.shape==(2766, 24, 1)
True
```
```
>>> np.all(Xtrain_e[:,:,0]==Xtrain) and np.all(ytrain_e[:,:,0]==ytrain) and np.all(Xtest_e[:,:,0]==Xtest) and np.all(ytest_e[:,:,0]==ytest)
True
```


Tests for q8 (points: None)
```
>>> model_dense.get_config()['layers'][0]['class_name']=='Flatten' and \
... model_dense.get_config()['layers'][1]['class_name']=='Dense' and \
... model_dense.get_config()['layers'][1]['config']['units']==5 and \
... model_dense.get_config()['layers'][1]['config']['activation']=='relu' and \
... model_dense.get_config()['layers'][2]['class_name'] and \
... model_dense.get_config()['layers'][2]['config']['units']==24 and \
... model_dense.get_config()['layers'][2]['config']['activation']=='linear'
True
```


Tests for q9 (points: None)
```
>>> model_simprnn.get_config()['layers'][0]['class_name']=='InputLayer' and \
... model_simprnn.get_config()['layers'][0]['config']['batch_input_shape']==(None, 6, 1) and \
... model_simprnn.get_config()['layers'][1]['class_name']=='SimpleRNN' and \
... model_simprnn.get_config()['layers'][1]['config']['batch_input_shape']==(None, 6, 1) and \
... model_simprnn.get_config()['layers'][2]['class_name']=='Dense' and \
... model_simprnn.get_config()['layers'][2]['config']['units']==24 
True
```
```
>>> np.all(np.isclose(model_simprnn.get_weights()[0],np.array([[-1.1073184 ,  0.95952076, -0.7785383 ,  0.77206063, -0.00925372]]),1e-3))
False
```
