```python
# Initialize Otter
import otter
grader = otter.Notebook("hw3.ipynb")
```

---

<h1><center>SDSE Homework 3 <br><br> Simple linear regression </center></h1>

---

```python
import numpy as np
import scipy.stats as stats
import matplotlib.pyplot as plt
import pickle
```

# 0. Load the data

In this homework we will implement simple linear regression and apply it to four different datasets. 

All of the formulas can be found in the lecture and reader. 

The four datasets are contained in `hw3data.pickle`. Each dataset is a tuple with entries `x` and `y`:  `DA = (x,y)`.

```python
with open('hw3data.pickle','rb') as f:
    DA, DB, DC, DD = pickle.load(f)
```

# 1. Scatter plots

Display each of the four datasets using matplotlib [scatter](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.scatter.html) plots. Your plot should look similar to the one below, with two rows and two columns of plots.  

<img src="f1.png" />

```python
fig, ax = plt.subplots(nrows=...,ncols=...)  # TODO
ax[...,...].scatter(...,...,marker='.')
...
```

Output:
```
File ~/venvs/sdseenv/lib/python3.10/site-packages/matplotlib/gridspec.py:379, in GridSpec.__init__(self, nrows, ncols, figure, left, bottom, right, top, wspace, hspace, width_ratios, height_ratios)
    376 self.hspace = hspace
    377 self.figure = figure
--> 379 super().__init__(nrows, ncols,
    380                  width_ratios=width_ratios,
    381                  height_ratios=height_ratios)

File ~/venvs/sdseenv/lib/python3.10/site-packages/matplotlib/gridspec.py:49, in GridSpecBase.__init__(self, nrows, ncols, height_ratios, width_ratios)
     34 """
     35 Parameters
     36 ----------
   (...)
     46     If not given, all rows will have the same height.
     47 """
     48 if not isinstance(nrows, Integral) or nrows <= 0:
---> 49     raise ValueError(
     50         f"Number of rows must be a positive integer, not {nrows!r}")
     51 if not isinstance(ncols, Integral) or ncols <= 0:
     52     raise ValueError(
     53         f"Number of columns must be a positive intege
```

# 2. Sample correlation coefficient

Implement the function `corr(D)` that takes a dataset tuple `D` and returns its sample correlation coefficient $r$. 

```python
def corr(D):
    x, y = D
    ...             # TODO
    return ...      # TODO
```

```python
rA = corr(DA)
rB = corr(DB)
rC = corr(DC)
rD = corr(DD)

print(f'rA={rA}')
print(f'rB={rB}')
print(f'rC={rC}')
print(f'rD={rD}')
```

```python
grader.check("q2")
```

# 3. Compute linear regresssion coefficients

Implement the function `compute_linear_regression(D)` that takes a dataset tuple `D` and returns the coefficients of the linear regression $\theta_0$ and $\theta_1$. 

```python
def compute_linear_regression(D):
    ...                 # TODO
    theta1 = ...        # TODO
    theta0 = ...        # TODO
    return theta0, theta1
```

```python
thetaA = compute_linear_regression(DA)
thetaB = compute_linear_regression(DB)
thetaC = compute_linear_regression(DC)
thetaD = compute_linear_regression(DD)

print(f'A: {thetaA}')
print(f'B: {thetaB}')
print(f'C: {thetaC}')
print(f'D: {thetaD}')
```

```python
grader.check("q3")
```

# 4. Evaluate linear regresssion prediction

Implement the function `eval_linreg(theta,D)` that takes parameters $\theta=(\theta_0,\theta_1$ and the dataset `D=(x,y)` and returns the array of predictions $\hat{y}$ for the training data: 
$$\hat{y} = \theta_0 + x \theta_1$$

```python
def eval_linreg(theta,D):
    ...             # TODO
    return ...      # TODO
```

```python
yhatA = eval_linreg(thetaA,DA)
yhatB = eval_linreg(thetaB,DB)
yhatC = eval_linreg(thetaC,DC)
yhatD = eval_linreg(thetaD,DD)

print(f'yhatA[:20:3]={yhatA[:20:3]}')
print(f'yhatB[:20:3]={yhatB[:20:3]}')
print(f'yhatC[:20:3]={yhatC[:20:3]}')
print(f'yhatD[:20:3]={yhatD[:20:3]}')
```

```python
grader.check("q4")
```

# 5. Plot linear regresssion prediction

Overlay the linear regression lines on each one of the scatter plots. Your plot should look like this:

<img src="f2.png" />

```python
fig, ax = plt.subplots(figsize=(12,8),nrows=2,ncols=2)
ax[0,0].scatter(...,...,color='k',marker='.')
ax[0,0].plot(...,...)
...
```

# 6. Coefficient of determinination

Implement the function `eval_Rsquared(theta,D)` that takes parameters $\theta=(\theta_0,\theta_1)$ and the dataset `D` and returns the coefficient of determination $R^2$. Verify the identity $R^2=r^2$


```python
def eval_Rsquared(theta,D):
    ...             # TODO
    return ...      # TODO
```

```python
RsquaredA = eval_Rsquared(thetaA,DA)
RsquaredB = eval_Rsquared(thetaB,DB)
RsquaredC = eval_Rsquared(thetaC,DC)
RsquaredD = eval_Rsquared(thetaD,DD)

print(f'RsquaredA={RsquaredA}')
print(f'RsquaredB={RsquaredB}')
print(f'RsquaredC={RsquaredC}')
print(f'RsquaredD={RsquaredD}')
```

# 7. Estimate $\sigma^2$

Implement the function `estimate_sigma2(theta,D)` that takes parameters $\theta=(\theta_0,\theta_1)$ and the dataset `D` and returns an unbiased estimate of the variance $\sigma^2$ of the uncertainty values $\varepsilon_i$ in $y_i=\theta_0^* + x_i\theta_1^*  + \varepsilon_i$. 

$$\hat\sigma^2  = \frac{1}{N-2} \sum_{i=1}^{N} (y_i - \hat{y}_i)^2$$


```python
def estimate_sigma2(theta,D):
    ...             # TODO
    return ...      # TODO
```

```python
sigma2A = estimate_sigma2(thetaA,DA)
sigma2B = estimate_sigma2(thetaB,DB)
sigma2C = estimate_sigma2(thetaC,DC)
sigma2D = estimate_sigma2(thetaD,DD)

print(f'sigma2A={sigma2A}')
print(f'sigma2B={sigma2B}')
print(f'sigma2C={sigma2C}')
print(f'sigma2D={sigma2D}')
```

```python
grader.check("q7")
```

# 8. Uncertainty in the linear regression coefficients

Implement the function `estimate_sigma2(theta,D)` that takes parameters $\theta=(\theta_0,\theta_1)$ and the dataset `D` and returns the variance associated with the linear regression parameter estimates $Var[\hat\Theta_0]$ and $Var[\hat\Theta_1]$.

```python
def compute_var_param(theta,D):
    ...             # TODO
    var0 = ...      # TODO
    var1 = ...      # TODO
    return var0, var1
```

```python
DC
```

```python
varA = compute_var_param(thetaA,DA)
varB = compute_var_param(thetaB,DB)
varC = compute_var_param(thetaC,DC)
varD = compute_var_param(thetaD,DD)

print(f'A: {varA}')
print(f'B: {varB}')
print(f'C: {varC}')
print(f'D: {varD}')
```

```python
grader.check("q8")
```

# 9. Confidence intervals and hypothesis tests

Implement the function `compute_rho(theta,D,gamma)` that takes parameters $\theta=(\theta_0,\theta_1)$, the dataset `D`, and a confidence level $\gamma$, and returns the radiuses $\rho_0$ and $\rho_1$ of the $\gamma$-confidence intervals for $\theta^*_0$ and $\theta^*_1$ respectively. 

```python
def compute_rho(theta,D,gamma):
    ...             # TODO
    rho0 = ...      # TODO
    rho1 = ...      # TODO
    return rho0, rho1
```

```python
gamma = 0.98
rhoA = compute_rho(thetaA,DA,gamma)
rhoB = compute_rho(thetaB,DB,gamma)
rhoC = compute_rho(thetaC,DC,gamma)
rhoD = compute_rho(thetaD,DD,gamma)

print(f'rhoA={rhoA}')
print(f'rhoB={rhoB}')
print(f'rhoC={rhoC}')
print(f'rhoD={rhoD}')
```

```python
grader.check("q9")
```

# 10. Negligible parameters
For which of the slope parameters `thetaA[1]`, `thetaB[1]`, `thetaC[1]`, or `thetaD[1]` is the null hypothesis 
$$ H_0 : \theta^*_1=0$$
**not** rejected in favor of
$$ H_1 : \theta^*_1\neq0$$
with $\alpha=0.05$ significance level? 

Provide your answer as a string from `{'A','B','C','D'}`

```python
answer = ...      # TODO
```

# 11. Nonlinear feature vector

The dataset `DA` looks like it could be modeled as an exponential decay. Let's use this insight to see whether we can improve the regression model for this dataset. Define a feture vector `phi` with
$$\phi(x) = e^{-\lambda x}$$
Using trial-and-error, find a value for $\lambda$ that achieves an $R^2$ of 0.9 or more. 

```python
lbda = ...              # TODO
x, y = DA
phi = np.exp(lbda*x)
DA2 = (phi,y)
thetaA2 = ...           # TODO
RsquaredA2 = ...        # TODO

fig, ax = plt.subplots()
ax.scatter(DA[0],DA[1],color='k',marker='.')
ax.plot(DA[0],eval_linreg(thetaA,DA),label='$R^2$={:.3f}'.format(RsquaredA))
ax.plot(DA[0],eval_linreg(thetaA2,DA2),'-',label='$R^2$={:.3f}'.format(RsquaredA2))
ax.legend(fontsize=14)
```

```python
grader.check("q11")
```

# 12. Prediction uncertainty


Implement the function `compute_var_prediction(theta,D,gamma,xi)` that takes parameters $\theta$, the dataset `D`, and a data sample $x_i$, and returns the variance in the prediction $\hat{y}_i$ corresponding to $x_i$.

```python
def compute_var_prediction(theta,D,xi):
    ...             # TODO
    return ...      # TODO
```

```python
varYA = compute_var_prediction(thetaA,DA,DA[0])
stddevYA = np.sqrt(varYA)

varYB = compute_var_prediction(thetaB,DB,DB[0])
stddevYB = np.sqrt(varYB)

varYC = compute_var_prediction(thetaC,DC,DC[0])
stddevYC = np.sqrt(varYC)

varYD = compute_var_prediction(thetaD,DD,DD[0])
stddevYD = np.sqrt(varYD)

print(f'stddevYA={stddevYA[:20:3]}')
print(f'stddevYB={stddevYB[:20:3]}')
print(f'stddevYC={stddevYC[:20:3]}')
print(f'stddevYD={stddevYD[:20:3]}')
```

```python
grader.check("q12")
```

# 13. Plot the prediction with uncertainty bounds

Repeat the scatter and prediction plot, but this time overlay an prediction uncertainty band of plus/minus two standard deviations. Your plot should look something like this:

<img src="f3.png" />


```python
fig, ax = plt.subplots(figsize=(12,8),nrows=2,ncols=2)

a = ax[0,0]
yhatA = eval_linreg(thetaA,DA)
a.scatter(DA[0],DA[1],color='k',marker='.')
a.plot(DA[0],yhatA)
a.fill_between(DA[0],yhatA-2*stddevYA,yhatA+2*stddevYA,alpha=0.5)

...
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q10 (points: None)


Tests for q11 (points: None)
```
>>> phi = np.exp(lbda*DA[0])
>>> DA2 = (phi,DA[1])
>>> thetaA2 = compute_linear_regression(DA2)
>>> RsquaredA2 = eval_Rsquared(thetaA2,DA2)
>>> RsquaredA2>0.9
True
```


Tests for q12 (points: None)
```
>>> np.all(np.isclose(stddevYA[:20:3],[0.08643448, 0.06966488, 0.0498225,  0.04820286, 0.04679432, 0.04424612,  0.0414014 ],1e-3))
True
```
```
>>> np.all(np.isclose(stddevYB[:20:3],[0.08762327, 0.08352081, 0.08168098, 0.07779715, 0.07289334, 0.07035215,  0.06925625],1e-3))
True
```
```
>>> np.all(np.isclose(stddevYC[:20:3],[4.35776707, 3.33866243, 2.85387726, 2.42265468, 2.23691524, 2.12253484,  1.95435833],1e-3))
True
```
```
>>> np.all(np.isclose(stddevYD[:20:3],[1.05420215, 0.72793851, 0.65995923, 0.55575209, 0.49422517, 0.47528383,  0.43187652],1e-3))
True
```


Tests for q2 (points: None)
```
>>> np.isclose(rA,-0.8099147715792736,1e-3)
True
```


Tests for q3 (points: None)
```
>>> np.isclose(thetaA[0],3.749057503898907,1e-3) and np.isclose(thetaA[1],-0.0462178161454909,1e-3)
True
```


Tests for q4 (points: None)
```
>>> np.all(np.isclose(yhatA[:20:3],[3.98923839, 3.73908805, 3.4193214, 3.39071521, 3.36528081, 3.31764243, 3.26114505],1e-3))
True
```
```
>>> np.all(np.isclose(yhatB[:20:3],[1.10744508, 1.2380595, 1.29750078, 1.42506595, 1.59110387, 1.67990742, 1.71890906],1e-3))
True
```
```
>>> np.all(np.isclose(yhatC[:20:3],[179.72584278, 156.97516769, 145.60977061, 134.87236202, 129.92229685, 126.71703594, 121.67884977],1e-2))
True
```
```
>>> np.all(np.isclose(yhatD[:20:3],[50.46531219, 50.17274557, 50.10828722, 50.00413843, 49.93722651, 49.91523581, 49.86076567],1e-2))
True
```


Tests for q6 (points: None)


Tests for q7 (points: None)
```
>>> np.isclose(sigma2A,0.0962882921462249,1e-3)
True
```
```
>>> np.isclose(sigma2B,0.2082814866175884,1e-3)
True
```
```
>>> np.isclose(sigma2C,218.80072616698286,1e-1)
True
```
```
>>> np.isclose(sigma2D,11.604657865691147,1e-1)
True
```


Tests for q8 (points: None)
```
>>> np.isclose(varA[0],0.004944706915203485,1e-3) and np.isclose(varA[1],1.1431957976526583e-05,1e-3) 
True
```
```
>>> np.isclose(varB[0],0.008480394564633838,1e-5) and np.isclose(varB[1],6.201746581241882e-05,1e-2) 
True
```
```
>>> np.isclose(varC[0],4.5691174071390455,1e-1) and np.isclose(varC[1],0.005635582130422289,1e-2)
True
```
```
>>> np.isclose(varD[0],0.21272847310590542,1e-3) and np.isclose(varD[1],0.00025779124994615524,1e-5) 
True
```


Tests for q9 (points: None)
```
>>> np.isclose(rhoA[0],0.16358554892650395,1e-3) and np.isclose(rhoA[1],0.007865656343192947,1e-5)
True
```
```
>>> np.isclose(rhoB[0],0.214231183811987,1e-3) and np.isclose(rhoB[1],0.018320261402565527,1e-4)
True
```
```
>>> np.isclose(rhoC[0],4.97268348585844,1e-2) and np.isclose(rhoC[1],0.1746401315487988,1e-3)
True
```
```
>>> np.isclose(rhoD[0],1.072969728803447,1e-2) and np.isclose(rhoD[1],0.037351559940089984,1e-4)
True
```
